# Lab 01: CSS 1

Welcome to the first lab of **CSS 1**!

In this self-paced lab, you'll get hands-on practice with the following **tools** and **concepts**:

- Jupyter notebooks.  
- Defining and using variables.  
- Submitting assignments using DataHub.

### Guide to CSS 1 Labs

Programming can be challenging, whether you're experienced or completely new. But it can also be a lot of fun. If you find yourself overwhelmed or frustrated with a question, remember: we're here to help!

- If you're working on this in section, feel free to ask your TA for guidance.  
- Labs are also **collaborative**, so you can also ask your fellow classmates. In fact, we encourage it!
- If you're stuck, you can also post questions on Piazza.

## Part 1: Jupyter notebook tooling

> **Tooling** means setting up all the software and tools you'll be using.

In this class, we're going to be working *a lot* with Jupyter notebooks, so Part 1 of this lab will focus on familiarizing you with how a Jupyter notebook works.

### Working with "cells"

> A **cell** is an independent unit in a Jupyter notebook.

This cell is a *Markdown* cell, which means you can type text into it along with code to *format* that text. Other cells are *Code* cells, which means you can use them to execute Python code.

You can determine the type of a cell using the dropdown menu on the Jupyter toolbar.

### Markdown Cells

#### Editing a Markdown cell.

If you double-click a cell in Jupyter notebooks, you can *edit* that cell. Try double-clicking the cell just below this one, and write: "My name is NAME" (where "NAME" is replaced by your own name).

<!-- grade_id: cell-449047da9a748b1b, points: 1, grade: true, solution: true -->


#### Using Markdown formatting.

You can also format the text in Markdown cells. A single asterisk `*` makes text *italic*, and two asterisks `**` makes text **bolded**. Use the single asterisk to write your *Intended Major* in the cell below.

<!-- grade_id: cell-bd654543a23bf29e, points: 1, grade: true, solution: true -->


#### Adding a new cell.

You can also create *new* cells. In the toolbar, you should see a "+" sign, which can use to create a new cell.

To create a new **Markdown** cell just below this one:

1. First, single-click on this cell (the one you're reading right now).
2. Then, press the "+" button.  
3. In that cell, use the dropdown toolbar depicted earlier to change the cell type to **Markdown**.
4. Then, complete the following sentence: "I'm taking CSS because..."

### Code cells

The other main type of Jupyter notebook cell we'll be using is a **code cell**. These can be used to write and run Python code.

#### Running a code cell

Just below this cell is a **code cell**. You can write Python code in that cell.

To run that cell, follow these instructions:

1. First, remove the code already there, where it says: `raise NotImplementedError`
2. Then, write the following code in the cell: `print("Hello!")`.
3. Press the "play/run" button in the toolbar.  
4. Alternatively, press "Shift + Enter".

<!-- grade_id: cell-505b3d7277c68522, points: 1, grade: true, solution: true -->
```python
print("Hello!")
```

Output:
```
Hello!
```

#### Modifying another code cell

The cell below this one will `raise` an **error** (also called an `Exception`) if you run it––or, notably, if you submit to DataHub without changing it.

Replace (i.e., delete and write something else) the code below with another `print` statement; make sure you run the cell to test if it works.

<!-- grade_id: cell-3b0426ecd747db9f, points: 0, grade: false, solution: true -->
```python
print("Hello, my name is Donovan Le!")
```

Output:
```
Hello, my name is Donovan Le!
```

## Part 2: Python basics

In this part, we'll get some hands-on practice with [Python basics](https://ucsd-css1-introduction.github.io/overview/lectures/03-basics-variables.html): declaring and using **variables**, using basic **operators** (e.g., `+` and `-`), and so on.

### Declaring variables

> A **variable** stores a particular value, like a "container".

Variables can be **declared** using the assignment operator (`=`).

#### Declaring an `int`

In the code cell below, create a variable called `my_first_int` and assign it to some number (e.g., `5`).

<!-- grade_id: cell-d115018370f00d1f, points: 0, grade: false, solution: true -->
```python
my_first_int = 1
my_first_int
```

Output:
```
1
```

Below is a piece of code that *checks* whether you've correctly declared this variable. If you haven't declared it yet, that cell will throw an error.

<!-- grade_id: cell-078d0746aca39b71, points: 1, grade: true, solution: false -->
```python
assert my_first_int
```

Below this cell is a piece of code that *checks* whether this variable has been assigned to an `int` type (as opposed to a `str`, etc.). If you've incorrectly assigned `my_first_int` to a `str`, it'll throw an error.

<!-- grade_id: cell-140beb2dc935b143, points: 1, grade: true, solution: false -->
```python
assert type(my_first_int) == int
```

#### Declaring a `str`

Now let's create another variable called `my_first_str`. Make sure you assign it to a `str` type, i.e., a sequence of characters surrounded quotes (e.g., `"apple"`).

<!-- grade_id: cell-a8877302b82f6070, points: 0, grade: false, solution: true -->
```python
my_first_str = "ts pmo gng"
my_first_str
```

Output:
```
'ts pmo gng'
```

Below is a piece of code that *checks* whether you've correctly declared this variable. If you haven't declared it yet, that cell will throw an error.

<!-- grade_id: cell-5d096511c1caa9e0, points: 1, grade: true, solution: false -->
```python
assert my_first_str
```

Below this cell is a piece of code that *checks* whether this variable has been assigned to a `str` type (as opposed to a `bool`, etc.). If you've incorrectly assigned `my_first_str` to a `bool` or some other type, it'll throw an error.

<!-- grade_id: cell-764de1ed9f1cefc4, points: 1, grade: true, solution: false -->
```python
assert type(my_first_str) == str
```

#### Declaring another `str`

In the cell below, create another `str` called `my_second_str`. The only constraint is that this `my_second_str` should have *more characters* than `my_first_str`.

<!-- grade_id: cell-e31c231b582d75ae, points: 0, grade: false, solution: true -->
```python
my_second_str = "lol ngl gang ts pmo sybau"
my_second_str
```

Output:
```
'lol ngl gang ts pmo sybau'
```

<!-- grade_id: cell-a261b60f16b8f03f, points: 1, grade: true, solution: false -->
```python
assert type(my_second_str) == str
```

<!-- grade_id: cell-2843c80287b9c2b5, points: 1, grade: true, solution: false -->
```python
assert my_second_str
```

<!-- grade_id: cell-4bcbca944bb911db, points: 1, grade: true, solution: false -->
```python
assert len(my_second_str) > len(my_first_str)
```

## Part 3: Submit

Once you've completed all the cells above (saving regularly):

- You will need to save your file on Google Drive
- Go to Canvas, to the Vocareum tab
- Select the appropriate assignment on Vocareum
- Submit your completed assignment
- Click "Validate"; Vocareum will tell you if you passed the public tests
- Submit your assignment on Vocareum

If you have any trouble accessing or submitting the assignment, please check in with your TA!